<img src="Archivos/miad4.png" width=800x>

# Modificar, limpiar y unir datos utilizando `pandas`

En este taller modificarás, limpiarás y unirás bases de datos al resolver diversos ejercicios.

## Habilidades en práctica

Al realizar este taller podrás revisar tu progreso para:

**1.** Modificar cadenas de texto contenidas en un `Series`.<br>
**2.** Representar fechas utilizando el paquete `datetime`.<br>
**3.** Imputar datos faltantes en un `DataFrame`.<br>
**4.** Unir dos `DataFrame` por coincidencia de filas o columnas.

## Instrucciones

En cada uno de los siguientes ejercicios deberás escribir el código solicitado estrictamente en las celdas indicadas para ello, teniendo en cuenta las siguientes recomendaciones:

* No crear, eliminar o modificar celdas de este Notebook (salvo lo que se te indique), pues puede verse afectado el proceso de calificación automática.

* La calificación se realiza de manera automática con datos diferentes a los proporcionados en este taller. Por consiguiente, tu código debe funcionar para diferentes instancias de cada uno de los ejercicios; una instancia hace referencia a los posibles valores de los parámetros.

* La calificación de cada ejercicio depende del valor que retorne la función especificada en su enunciado. Por lo tanto, aunque implementes funciones adicionales, es escencial que utilices los nombres propuestos en los enunciados de los ejercicios para implementar la función definitiva.

## Ejercicios
En la siguente celda encuentras declarados los paquetes necesarios para el desarrollo de este taller.

```python
# No modifiques esta celda
import numpy as np
import pandas as pd
from datetime import datetime
```

### Ejercicio 1

En la siguiente celda encuentras declarados dos `DataFrame` que almacenan los datos de los clientes de un negocio.

```python
# No modifiques esta celda
orden_cliente = pd.DataFrame([[70001, 3002],
                              [70002, 3001],
                              [70003, 3001],
                              [70004, 3003],
                              [70005, 3003]],
                              columns = ['no_orden', 'id_cliente'])

cliente_categoria = pd.DataFrame([[3001, "Platino"],
                                  [3002, "Oro"    ],
                                  [3003, "Oro"    ],
                                  [3004, "Bronce" ]],
                                  columns = ['id_cliente', 'Categoría Cliente'])
```

Implementa una función llamada `union_derecha` que reciba por parámetro dos `DataFrame` como los declarados en la celda anterior y que retorne un `DataFrame` resultado de la unión, donde se preserven todas las filas del segundo `DataFrame` y solo las que coincidan del primer `DataFrame`. La coincidencia debe hacerse evaluando la columna "id_cliente".

La función debe retornar un `DataFrame`.

<!-- grade_id: ej1_sol, points: 0, grade: false, solution: true -->
```python
## Implementa tu respuesta en esta celda
def union_derecha(dataframe1,dataframe2):
    union=pd.merge(dataframe1,dataframe2,how='right',on='id_cliente')
    return union


```

<!-- grade_id: ej1_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
orden_cliente = pd.DataFrame([[70001, 3002],
                              [70002, 3001],
                              [70003, 3001],
                              [70004, 3003],
                              [70005, 3003]],
                              columns = ['no_orden', 'id_cliente'])

cliente_categoria = pd.DataFrame([[3001, "Platino"],
                                  [3002, "Oro"    ],
                                  [3003, "Oro"    ],
                                  [3004, "Bronce" ]],
                                  columns = ['id_cliente', 'Categoría Cliente'])

orden_cliente_prueba= pd.DataFrame([[70001, 3002],
                                    [70002, 3002],
                                    [70003, 3001],
                                    [70004, 3000],
                                    [70006, 3004]],
                                   columns = ['no_orden', 'id_cliente'])

cliente_categoria_prueba = pd.DataFrame([[3001, "Platino"],
                                         [3002, "Oro"    ],
                                         [3000, "Oro"    ],
                                         [3001, "Bronce" ]],
                                        columns = ['id_cliente', 'Categoría Cliente'])

# Caso 1: no existe la función.
try:
    union_derecha
    assert type(union_derecha) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada union_derecha.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    union_derecha(orden_cliente, cliente_categoria)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame.
assert type(union_derecha(orden_cliente, cliente_categoria)) == type(orden_cliente), f"Tu función debe retornar un objeto de tipo '{type(orden_cliente).__name__}'."

# Caso 4: retorna un dataframe con cantidad de columnas errada
assert union_derecha(orden_cliente, cliente_categoria).shape[1] == 3, "Tu función retorna un DataFrame con cantidad de columnas errada."
assert union_derecha(orden_cliente_prueba, cliente_categoria_prueba).shape[1] == 3, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 5: devuelve un dataframe con cantidad de filas errada
assert union_derecha(orden_cliente, cliente_categoria).shape[0] == 6, "Tu función retorna un DataFrame con cantidad de filas errada."
assert union_derecha(orden_cliente_prueba, cliente_categoria_prueba).shape[0] == 5, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 6: retorna los valores que son (dfs prueba)
assert union_derecha(orden_cliente_prueba, cliente_categoria_prueba).iat[4,2] == 'Bronce', "Tu función retorna un DataFrame con valores distintos a los esperados."
assert union_derecha(orden_cliente_prueba, cliente_categoria_prueba).iat[4,1] == 3001, "Tu función retorna un DataFrame con valores distintos a los esperados."

### BEGIN HIDDEN TESTS
'''
from pandas.testing import assert_frame_equal

def union_derecha_sol(df1, df2):
    resultado = df1.merge(df2, how = 'right', left_on = "id_cliente", right_on = "id_cliente")        
    return resultado

students = union_derecha(orden_cliente, cliente_categoria)
ours = union_derecha_sol(orden_cliente, cliente_categoria)

try:
    assert_frame_equal(students, ours)
except AssertionError as e:
    e.args += ("Tu respuesta es incorrecta.",)
    raise
''' 
### END HIDDEN TESTS

print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 2

En la siguiente celda encuentras declarados dos `DataFrame` que almacenan los datos de los clientes de un negocio.

```python
# No modifiques esta celda
cliente_categoria_orden = pd.DataFrame([[3002, "Platino", 70001],
                                        [3004, "Oro",     70002],
                                        [3003, "Oro",     70003]],
                                        columns = ['id_cliente', 'Categoría Cliente', 'no_orden'])
    
cliente_categoria = pd.DataFrame([[3006, "Oro"    ],
                                  [3001, "Oro"    ],
                                  [3005, "Platino"],
                                  [3007, "Platino"]],
                                  columns = ['id_cliente', 'Categoría Cliente'])
```

Implementa una función llamada `concatenar_interno` que reciba por parámetro dos `DataFrame` como los declarados en la celda anterior y que retorne un `DataFrame` resultado de la unión, donde se preserven las filas de ambos `DataFrame`.

La función deber retornar un `DataFrame` con las **columnas** que existan en ambos `DataFrame` y todas las filas de ambos.

<!-- grade_id: ej2_sol, points: 0, grade: false, solution: true -->
```python

def concatenar_interno(dataframe1, dataframe2):
    # Encontramos las columnas que tienen en comun.
    common_columns = dataframe1.columns.intersection(dataframe2.columns).tolist()
    
    # Unimos las dos bases de datos.
    result = pd.merge(dataframe1[common_columns], dataframe2[common_columns], how='outer')
    return result



```

<!-- grade_id: ej2_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
cliente_categoria_orden = pd.DataFrame([[3002, "Platino", 70001],
                                        [3004, "Oro",     70002],
                                        [3003, "Oro",     70003]],
                                        columns = ['id_cliente', 'Categoría Cliente', 'no_orden'])
    
cliente_categoria = pd.DataFrame([[3006, "Oro"    ],
                                  [3001, "Oro"    ],
                                  [3005, "Platino"]],
                                  columns = ['id_cliente', 'Categoría Cliente'])

cliente_categoria_orden_prueba = pd.DataFrame([[3002, "Platino", 70001],
                                        [3004, "Oro",     70002],
                                        [3003, "Oro",     70003]],
                                        columns = ['id_cliente', 'Categoría Cliente', 'no_orden'])
    
cliente_categoria_prueba = pd.DataFrame([[3006, 70006],
                                  [3001, 70007],
                                  [3009, 70008]],
                                  columns = ['id_cliente', 'no_orden'])

# Caso 1: no existe la función.
try:
    concatenar_interno
    assert type(concatenar_interno) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada union_interna.",)
    
# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    concatenar_interno(cliente_categoria_orden, cliente_categoria)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame.
assert type(concatenar_interno(cliente_categoria_orden, cliente_categoria)) == type(cliente_categoria), f"Tu función debe retornar un objeto de tipo '{type(orden_cliente).__name__}'."

# Caso 4: retorna un dataframe con cantidad de columnas errada
assert concatenar_interno(cliente_categoria_orden, cliente_categoria).shape[1] == 2, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 5: devuelve un dataframe con cantidad de filas errada
assert concatenar_interno(cliente_categoria_orden, cliente_categoria).shape[0] == 6, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 7: devuelve un dataframe cuyas columnas tienen nombre distinto de lo esperado (dfs prueba)
assert [*concatenar_interno(cliente_categoria_orden_prueba, cliente_categoria_prueba).columns] == ['id_cliente','no_orden'], "Tu función retorna un DataFrame con las columnas incorrectas."


#Caso 8: retorna un dataframe con los valores correctos (dfs prueba)
assert concatenar_interno(cliente_categoria_orden_prueba, cliente_categoria_prueba).iat[1,1] == 70002, "Tu función retorna un DataFrame con los valores incorrectos."
assert concatenar_interno(cliente_categoria_orden_prueba, cliente_categoria_prueba).iat[5,0] == 3009,"Tu función retorna un DataFrame con los valores incorrectos."

### BEGIN HIDDEN TESTS
'''
def concatenar_interno_sol(df, df2):
    
    resultado = pd.concat([df,df2], axis = 0, join = "inner")
    
    return resultado

students = concatenar_interno(cliente_categoria_orden, cliente_categoria)
ours = concatenar_interno_sol(cliente_categoria_orden, cliente_categoria)
try:
    assert_frame_equal(students, ours)
except AssertionError as e:
    e.args += ('Tu respuesta es incorrecta',)
    raise
'''
### END HIDDEN TESTS
print("Felicidades, realizaste este ejercicio correctamente.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 3

En la siguiente celda encuentras declarado un `DataFrame` que almacena las ventas diarias de los tres principales productos de un negocio.

```python
# No modifiques esta celda
ventas = pd.DataFrame([[2600.0,  12200.0, 6000.0],
                       [11200.0, 9000.0,  600.0 ],
                       [np.nan,  8000.0,  3000.0],
                       [4000.0,  np.nan,  np.nan],
                       [np.nan,  10000.0, 2000.0]],
                      columns = ["Producto 1", "Producto 2", "Producto 3"])
```

Implementa una función llamada `mediana_faltantes` que reciba por parámetro un `DataFrame` como el declarado en la celda anterior, complete los datos faltantes del `DataFrame` con la mediana de su columna respectiva y lo retorne.

La función deber retornar un `DataFrame` sin datos faltantes.

<!-- grade_id: ej3_sol, points: 0, grade: false, solution: true -->
```python
## Implementa tu respuesta en esta celda
def mediana_faltantes(dataframe):
    mediana=dataframe.median(axis=0)
    dataframe=dataframe.fillna(mediana)
    return dataframe

mediana=mediana_faltantes(ventas)
print(mediana)
```

Output:
```
Producto 1  Producto 2  Producto 3
0      2600.0     12200.0      6000.0
1     11200.0      9000.0       600.0
2      4000.0      8000.0      3000.0
3      4000.0      9500.0      2500.0
4      4000.0     10000.0      2000.0
```

<!-- grade_id: ej3_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
ventas = pd.DataFrame([[2600.0,  12200.0, 6000.0],
                       [11200.0, 9000.0,  600.0 ],
                       [np.nan,  8000.0,  3000.0],
                       [4000.0,  np.nan,  np.nan],
                       [np.nan,  10000.0, 2000.0]],
                      columns = ["Producto 1", "Producto 2", "Producto 3"])

ventas_prueba = pd.DataFrame([[2600.0,  12200.0, 6000.0],
                       [11200.0, 9000.0,  6000.0 ],
                       [3000.0,  8000.0,  np.nan],
                       [np.nan,  np.nan,  6500.0],
                       [np.nan,  12000.0, 2000.0]],
                      columns = ["Producto 1", "Producto 2", "Producto 3"])

# Caso 1: no existe la función.
try:
    mediana_faltantes
    assert type(mediana_faltantes) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada mediana_faltantes.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    mediana_faltantes(ventas)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame.
assert type(mediana_faltantes(ventas)) == type(ventas), f"Tu función debe retornar un objeto de tipo '{type(ventas).__name__}'."

# Caso 4: retorna un dataframe con cantidad de columnas errada
assert mediana_faltantes(ventas).shape[1] == 3, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 5: devuelve un dataframe con cantidad de filas errada
assert mediana_faltantes(ventas).shape[0] == 5, "Tu función retorna un DataFrame con cantidad de filas errada."

# Caso 5: no hay valores na
assert mediana_faltantes(ventas).isnull().sum().sum() == 0, "Tu función retorna un DataFrame con valores faltantes."
assert mediana_faltantes(ventas_prueba).isnull().sum().sum() == 0, "Tu función retorna un DataFrame con valores faltantes."
# Caso 6: Los espacios vacios se llenan con la mediana
assert mediana_faltantes(ventas_prueba).iat[3,0] == 3000.0, "Tu función retorna un DataFrame que no llena los valores faltantes con la mediana."
assert mediana_faltantes(ventas_prueba).iat[2,2] == 6000.0, "Tu función retorna un DataFrame que no llena los valores faltantes con la mediana."
#respuestas 
try:
    assert mediana_faltantes(ventas).iat[2,0] == 4000
    assert mediana_faltantes(ventas).iat[3,1] == 9500
except AssertionError as e:
    e.args += ("Tu función retorna un DataFrame que no llena los valores faltantes con la mediana.",)
    raise e
    
### BEGIN HIDDEN TESTS
'''
def mediana_faltantes_sol(df):
    resultado = df.fillna(df.median())
    return resultado

students = mediana_faltantes(ventas)
ours = mediana_faltantes_sol(ventas)

try:
    assert_frame_equal(students, ours)
except AssertionError as e:
    e.args += ('Tu respuesta es incorrecta',)
    raise
'''  
### END HIDDEN TESTS

print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 4

En la siguiente celda encuentras declarados dos `DataFrame` que almacenan los datos de los clientes de un negocio.

```python
# No modifiques esta celda
clientes = pd.DataFrame([[3002, "Masculino", "No"],
                         [3001, "Femenino",  "No"],
                         [3003, "Femenino",  "Sí"]],
                        columns = ["id_cliente", "Sexo", "Cliente Premium"])
```

Implemente una función llamada `cliente_dummies` que reciba por parámetro un `DataFrame` como el declarado en la celda anterior, y retorne un `DataFrame` con la representación en variables *dummy* de las columnas `"Sexo"` y `"Cliente Premium"`.

* La variable `"Sexo"` debe ser representada por un 1 si su valor es `"Femenino"` y cero de lo contrario.
* La variable `"Cliente Premium"` debe ser representada por un 1 si su valor es `"No"` y cero de lo contrario.

La función deber retornar un `DataFrame` con la misma cantidad de filas y columnas, pero números binarios en lugar de los registros de sexo y si es cliente premium.

<!-- grade_id: ej4_sol, points: 0, grade: false, solution: true -->
```python
def cliente_dummies(dataframe):
    # Crea variables dummy para la columna "Sexo"
    dataframe['Sexo'] = dataframe['Sexo'].apply(lambda x: 1 if x == 'Femenino' else 0)
    
    # Crea variables dummy para la columna "Cliente Premium"
    dataframe['Cliente Premium'] = dataframe['Cliente Premium'].apply(lambda x: 1 if x == 'No' else 0)
    
    return dataframe


```

<!-- grade_id: ej4_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
clientes = pd.DataFrame([[3002, "Masculino", "No"],
                         [3001, "Femenino",  "No"],
                         [3003, "Femenino",  "Sí"]],
                        columns = ["id_cliente", "Sexo", "Cliente Premium"])

clientes_prueba = pd.DataFrame([[3002, "Masculino", "Sí"],
                         [3001, "Masculino",  "No"],
                         [3003, "Femenino",  "Sí"]],
                        columns = ["id_cliente", "Sexo", "Cliente Premium"])

# Caso 1: no existe la función.
try:
    cliente_dummies
    assert type(cliente_dummies) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada cliente_dummies.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    cliente_dummies(clientes)
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame.
assert type(cliente_dummies(clientes)) == type(clientes), f"Tu función debe retornar un objeto de tipo '{type(clientes).__name__}'."

# Caso 4: retorna un dataframe con cantidad de columnas errada
assert cliente_dummies(clientes).shape[1] == 3, "Tu función retorna un DataFrame con cantidad de columnas errada."

# Caso 5: devuelve un dataframe con cantidad de filas errada
assert cliente_dummies(clientes).shape[0] == 3, "Tu función retorna un DataFrame con cantidad de filas errada."

#Caso 6: Tu funcion no retorna los dummies que son (prueba)
assert cliente_dummies(clientes_prueba).iat[1,1] == 0,"Tu función retorna un DataFrame con los valores incorrectos."
assert cliente_dummies(clientes_prueba).iat[1,2] == 1,"Tu función retorna un DataFrame con los valores incorrectos."


### BEGIN HIDDEN TESTS
'''
def cliente_dummies_sol(df):
    
    resultado = df.copy()
    resultado['Sexo'] = resultado['Sexo'].str.get_dummies()
    resultado['Cliente Premium'] = resultado['Cliente Premium'].str.get_dummies()
    
    return resultado

students = cliente_dummies(clientes)
ours = cliente_dummies_sol(clientes)
try:
    assert_frame_equal(students, ours)
except AssertionError as e:
    e.args += ('Tu respuesta es incorrecta',)
    raise
'''
### END HIDDEN TESTS
print("Felicidades, realizaste este ejercicio correctamente.")

```

Output:
```
AssertionError                            Traceback (most recent call last)
c:\Users\David\Documents\Analytics\Herramientas_Computacionales\Talleres\Taller7\TALLER - Limpiar y unir bases de datos.ipynb Cell 25 line 3
     <a href='vscode-notebook-cell:/c%3A/Users/David/Documents/Analytics/Herramientas_Computacionales/Talleres/Taller7/TALLER%20-%20Limpiar%20y%20unir%20bases%20de%20datos.ipynb#X33sZmlsZQ%3D%3D?line=35'>36</a> #Caso 6: Tu funcion no retorna los dummies que son (prueba)
     <a href='vscode-notebook-cell:/c%3A/Users/David/Documents/Analytics/Herramientas_Computacionales/Talleres/Taller7/TALLER%20-%20Limpiar%20y%20unir%20bases%20de%20datos.ipynb#X33sZmlsZQ%3D%3D?line=36'>37</a> assert cliente_dummies(clientes_prueba).iat[1,1] == 0,"Tu función retorna un DataFrame con los valores incorrectos."
---> <a href='vscode-notebook-cell:/c%3A/Users/David/Documents/Analytics/Herramientas_Computacionales/Talleres/Taller7/TALLER%20-%20Limpiar%20y%20unir%20bases%20de%20datos.ipynb#X33sZm
```

### Ejercicio 5

En la siguiente celda encuentras declarada una fecha tipo cadena de caracteres.

```python
# No modifiques esta celda
fecha_string = "2020-10-December"
```

Implemente una función llamada `cadena_a_fecha` que reciba por parámetro una cadena como la declarada en la celda anterior y retorne un objeto de tipo `datetime` que represente la misma fecha.

La función deber retornar un objeto de tipo `datetime`.

<!-- grade_id: ej5_sol, points: 0, grade: false, solution: true -->
```python
## Implementa tu respuesta en esta celda
def cadena_a_fecha(cadena):
    fecha=datetime.strptime(cadena,'%Y-%d-%B')
    return fecha

```

<!-- grade_id: ej5_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables

fecha_string = "2020-10-December"

fecha_string_prueba = "2021-20-October"

# Caso 1: no existe la función.
try:
    cadena_a_fecha
    assert type(cadena_a_fecha) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada cadena_a_fecha.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    cadena_a_fecha
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un Datetime.
assert type(cadena_a_fecha(fecha_string)) == datetime, f"Tu función debe retornar un objeto de tipo '{datetime.__name__}'."

#Caso 4: retorna la fecha correcta
assert cadena_a_fecha(fecha_string_prueba) == datetime(2021, 10, 20, 0, 0), f"Tu respuesta es incorrecta."
assert cadena_a_fecha(fecha_string) == datetime(2020, 12, 10, 0, 0), f"Tu respuesta es incorrecta."

### BEGIN HIDDEN TESTS
'''
def cadena_a_fecha_sol(cadena):

    resultado = datetime.strptime(fecha_string,"%Y-%d-%B")
    
    return resultado

students = cadena_a_fecha(fecha_string)
ours = cadena_a_fecha_sol(fecha_string)
try:
    assert students == ours
except AssertionError as e:
    e.args += ('Tu respuesta es incorrecta.',)
    raise
''' 
### END HIDDEN TESTS

print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 6

Junto a este Notebook encontrarás dos archivos de Excel. Cada uno de estos archivos contiene una base de datos, ambas provistas por el Banco de la República.

La primera base de datos (`"1.1.TIB_Serie historica IQY.xlsx"`) corresponde a la tasa a la que se puede prestar o pedir prestado entre los bancos.

La segunda base de datos (`"1.2.TIP_Serie historica diaria IQY.xlsx"`) corresponde a la tasa de interés de referencia fijada por el Banco de la República.

Implemente una función llamada `union_izquierda` que no reciba parámetros, importe los datos de los archivos de Excel y retorne un `DataFrame` resultado de la unión, donde se preserven todas las filas de los datos de la Tasa Interbancaria y solo las que coincidan de la Tasa de Intervención Pública. La coincidencia debe hacerse evaluando la columna "Fecha".

La función debe retornar un `DataFrame`.

<!-- grade_id: ej6_sol, points: 0, grade: false, solution: true -->
```python
def union_izquierda():


    df1 = pd.read_excel('Archivos/1.1.TIB_Serie historica IQY.xlsx', sheet_name='Sheet1', header=7)
    df2 = pd.read_excel('Archivos/1.2.TIP_Serie historica diaria IQY.xlsx', header=7)
    df1['Fecha(dd/mm/aaaa)'] = pd.to_datetime(df1['Fecha(dd/mm/aaaa)'], format='%Y-%m-%d', errors='coerce')
    df2['Fecha (dd/mm/aaaa)']=pd.to_datetime(df2['Fecha (dd/mm/aaaa)'], format='%Y-%m-%d', errors='coerce')

    resultado = pd.merge(df1, df2, how='left', left_on='Fecha(dd/mm/aaaa)', right_on='Fecha (dd/mm/aaaa)')
    return resultado

```

<!-- grade_id: ej6_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Caso 1: no existe la función.
try:
    union_izquierda
    assert type(union_izquierda) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada union_izquierda.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    union_izquierda
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un Dataframe.
assert type(union_izquierda()) == pd.core.frame.DataFrame, f"Tu función debe retornar un objeto de tipo '{pd.core.frame.DataFrame.__name__}'."

### BEGIN HIDDEN TESTS
'''
def union_izquierda_sol():

    TIB = pd.read_excel("Archivos/1.1.TIB_Serie historica IQY.xlsx",
                        index_col=0, usecols = [0,1], skiprows =7, skipfooter = 5) 
    TIP = pd.read_excel("Archivos/1.2.TIP_Serie historica diaria IQY.xlsx",
                        index_col=0, usecols = [0,1], skiprows =7, skipfooter = 5)
    resultado = TIP.join(TIB, how= "left")
    
    return resultado

students = union_izquierda()
ours = union_izquierda_sol()
try:
    assert_frame_equal(students, ours)
except AssertionError as e:
    e.args += ('Tu respuesta es incorrecta, intentalo de nuevo',)
    raise
'''
### END HIDDEN TESTS

print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 7

Crea una función llamada `ejercicio7` que reciba un `DataFrame` y retorne en una lista el nombre de las columnas que tienen más de 1 dato faltante. Puedes usar el `DataFrame` de la celda de código para probar que tu función este obteniendo el resultado correctamente. Recuerda que pueden usarse otros `DataFrame` para poner a prueba tu función.

La función `ejercicio7` debe retornar una lista.

```python
# No modifiques esta celda
df = pd.DataFrame({'no_orden':    [70001,np.nan,70002,70004,np.nan,70005,np.nan,70010,70003,70012,np.nan,70013],
                   'monto US$':   [150.5,270.65,65.26,110.5,948.5,2400.6,5760,1983.43,2480.4,250.45, 75.29,3045.6],
                   'Fecha Orden': ['2012-10-05','2012-09-10',np.nan,'2012-08-17','2012-09-10','2012-07-27','2012-09-10','2012-10-10','2012-10-10','2012-06-27','2012-08-17','2012-04-25'],
                   'idCliente':   [3002,3001,3001,3003,3002,3001,3001,3004,3003,3002,3001,3001],
                   'idVendedor':  [5002,5003,5001,np.nan,5002,5001,5001,np.nan,5003,5002,5003,np.nan]})
```

<!-- grade_id: ej7_sol, points: 0, grade: false, solution: true -->
```python
## Implementa tu respuesta en esta celda
def ejercicio7(df):
    n_columns=[]
    for columns in df.columns:
        count_na=df[columns].isnull().sum()
        if count_na>1:
            n_columns.append(columns)
    return n_columns




```

<!-- grade_id: ej7_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables

df = pd.DataFrame({
    'no_orden':[70001,np.nan,70002,70004,np.nan,70005,np.nan,70010,70003,70012,np.nan,70013],
    'monto US$':[150.5,270.65,65.26,110.5,948.5,2400.6,5760,1983.43,2480.4,250.45, 75.29,3045.6],
    'Fecha Orden': ['2012-10-05','2012-09-10',np.nan,'2012-08-17','2012-09-10','2012-07-27','2012-09-10','2012-10-10','2012-10-10','2012-06-27','2012-08-17','2012-04-25'],
    'idCliente':[3002,3001,3001,3003,3002,3001,3001,3004,3003,3002,3001,3001],
    'idVendedor':[5002,5003,5001,np.nan,5002,5001,5001,np.nan,5003,5002,5003,np.nan]})

df_prueba = pd.DataFrame({
    'no_orden':[70001,np.nan,70002,70004,np.nan,70005,np.nan,70010,70003,70012,np.nan,70013],
    'monto US$':[150.5,270.65,65.26,110.5,948.5,2400.6,5760,1983.43,2480.4,250.45, 75.29,3045.6],
    'Fecha Orden': ['2012-10-05','2012-09-10',np.nan,'2012-08-17','2012-09-10','2012-07-27','2012-09-10','2012-10-10','2012-10-10','2012-06-27','2012-08-17','2012-04-25'],
    'idCliente':[3002,3001,3001,3003,3002,np.nan,3001,np.nan,3003,3002,3001,3001],
    'idVendedor':[5002,5003,5001,5003,5002,5001,5001,5001,5003,5002,5003,np.nan]})

# Caso 1: no existe la función.
try:
    ejercicio7
    assert type(ejercicio7) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada ejercicio7.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    ejercicio7
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un Dataframe.
assert type(ejercicio7(df)) == list, f"Tu función debe retornar un objeto de tipo '{list.__name__}'."

# Caso 4: Probar con otro Dataframe que sean esas dos.
assert ejercicio7(df_prueba) == ['no_orden', 'idCliente'], f"Tu respuesta es incorrecta."

### BEGIN HIDDEN TESTS
'''
#probar con la otra respuesta? sera buena idea botarles la repsuesta

def ejercicio7_sol(df): 
    cantidad_faltantes = df.isnull().sum()  
    return [*cantidad_faltantes[cantidad_faltantes > 1].index]

assert ejercicio7(df) == ejercicio7_sol(df)
assert ejercicio7(df_prueba) == ejercicio7_sol(df_prueba)
'''
### END HIDDEN TESTS

print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 8

Colombia es un país con un gran número de días feriados. En la carpeta de este taller encontrarás una base de datos (`"Festivos.xlsx"`) que contiene estas fechas. Esta base de datos organiza en columnas los festivos de cada año, desde el año 2001 hasta el año 2019. Todos los festivos se presentan con el mismo formato: 

    6 Enero Día de los Reyes Magos

Crea una función llamada `ejercicio8` que procese la base de datos de festivos, de manera que retorne un objeto de tipo `DatetimeIndex` con todos los festivos que se han presentado en Colombia desde el año 2012 hasta el año 2020. 

Al imprimirse, el objeto tendría el siguientes aspecto:

    DatetimeIndex(['2001-01-01', '2001-01-08', '2001-03-19', '2001-04-12',
                   '2001-04-13', '2001-05-01', '2001-05-28', '2001-06-18',
                   '2001-06-25', '2001-07-02',
                   ...
                   '2020-06-29', '2020-07-20', '2020-08-07', '2020-08-17',
                   '2020-10-12', '2020-11-02', '2020-11-16', '2020-12-08',
                   '2020-12-25', '2020-12-31'],
                   dtype = 'datetime64[ns]', name = 'Fecha', length = 380, freq = None)
    
En la caja de código se le incluirá una lista con los meses del año. Esta lista le puede resultar util para procesar la información del `DataFrame`.

La función `ejercicio8` debe retornar un objeto `DatetimeIndex`.

```python
# No modifiques esta celda
lista_meses = ["Enero",      "Febrero", "Marzo",     "Abril",
               "Mayo",       "Junio",   "Julio",     "Agosto",
               "Septiembre", "Octubre", "Noviembre", "Diciembre"] 
```

<!-- grade_id: ej8_sol, points: 0, grade: false, solution: true -->
```python
## Implementa tu respuesta en esta celda

```

<!-- grade_id: ej8_test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
lista_meses = ["Enero", "Febrero", "Marzo", "Abril", "Mayo", "Junio", "Julio", "Agosto", "Septiembre",
               "Octubre", "Noviembre", "Diciembre"] 

# Caso 1: no existe la función.
try:
    ejercicio8
    assert type(ejercicio8) == type(lambda:None)
except:
    raise NotImplementedError("No existe una función llamada ejercicio8.",)

# Caso 2: la función es interrumpida por errores durante su ejecución.
try:
    ejercicio8
except:
    raise RuntimeError("Tu función produce un error al ejecutarse.")

# Caso 3: no retorna un Datatimeindex.
assert type(ejercicio8()) == pd.core.indexes.datetimes.DatetimeIndex, f"Tu función debe retornar un objeto de tipo '{pd.core.indexes.datetimes.DatetimeIndex.__name__}'."

# Caso 4: retorna un dataframe con cantidad de columnas errada
assert len(ejercicio8()) == 171, "Tu función retorna un objeto DatetimeIndex con una longitud incorrecta."

### BEGIN HIDDEN TESTS
lista_meses_ingles_sol = ["January",   "February", "March",    "April",
                          "May",       "June",     "July",     "August",
                          "September", "October",  "November", "December"] 

def ejercicio8_sol():    
    df = pd.read_excel('./Archivos/Festivos.xlsx')
    anhos = list(range(2012, 2021))
    df_filtr = df.loc[:, anhos]
    lista_fechas = []
    
    for columna in df_filtr.columns:        
        for indice in df_filtr.index:

            cadena = df.loc[indice, columna]
            cadena = cadena.replace(u'\xa0', u' ')
            dia = cadena.split(" ")[0]
            mes = lista_meses_ingles_sol[lista_meses.index(cadena.split(" ")[1])]
            
            fecha_string = str(columna) + "-" + mes + "-" + str(dia)
            fecha = datetime.strptime(fecha_string, "%Y-%B-%d")
            lista_fechas.append(fecha)
    
    return pd.DatetimeIndex(lista_fechas)


assert (ejercicio8() == ejercicio8_sol()).all

### END HIDDEN TESTS
print("Felicidades, realizaste este ejercicio correctamente.")

```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

## Referencias

Banco de la República (2020). 1.2.TIP_Serie histórica diaria IQY.xlsx [Base deDatos]. Recuperado el 15 de Noviembre de 2020 de :
https://www.banrep.gov.co/es/estadisticas/tasas-interes-politica-monetaria

Banco de la República (2020). 1.1.TIB_Serie histórica IQY.xlsx [Base deDatos]. Recuperado el 15 de Noviembre de 2020 de :
https://www.banrep.gov.co/es/estadisticas/tasa-interbancaria-tib

## Créditos

**Autor(es)**: Jorge Esteban Camargo Forero, Juan David Reyes Jaimes, Alejandro Mantilla Redondo, Diego Alejandro Cely Gomez

**Última actualización:** 17/09/2021